# <div style="background-color:rgba(204, 229, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(0, 76, 153, 1);">**PHYS 121 Pre-Lab #1**<span style="color:red"> $\to$ (9 possible marks)</span></span></div>

<font size ="4"><span style="color:purple">We recommend that you complete this notebook using either the [Google Chrome](https://www.google.com/intl/en_ca/chrome/) or [Mozilla Firefox](https://www.mozilla.org/en-CA/firefox/new/) browser.  Chrome and Firefox are both available for Windows, Mac, and Linux operating systems.  The PHYS 121 Jupyter notebooks should work in other browsers, but we have specifically verfied that they work well in both Chrome and Firefox.</span></font>

# Oscillations of a Simple Pendulum &ndash; Week 1

***
## Learning Objectives:
* <b><span style="color:rgba(0, 153, 76, 1);">Determine what *does* and *does not* affect the width of a distribution of measurements.</span></b>
* <b><span style="color:rgba(0, 153, 76, 1);">Determine the uncertainty associated with the *mean* or *average* value of a set of measurements.</span></b> 

***
## Autograding:
The PHYS 121 Pre-lab assignments and Labs will make use of some autograding.  To make the autograding work, the two cells below needs to be executed.  The first one installs the required packages and the second imports required packages/modules.  If 'PHYS121.Installer()' reports that some functions have been installed, the user should exectute the PHYS121.Installer() cell a second time.  The second time the installer function is run, it should report that **"All packages already installed. Please proceed"**.

If necessary, the kernel can be restarted by selecting **Kernel** $\to$ **Restart Kernel** from the menu options at to the top of the page.  Here is a <a href = "https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/general/gifs/restartKernel.gif">GIF</a> showing how to restart the kernel.

The 'PHYS121.Installer()' command requires the file 'PHYS121.py', which you should see included in the list of files along the left-hand side of the screen.

```python
# Hide unnecessary warnings.
import warnings
warnings.filterwarnings('ignore')

# Import PHYS121.py and then run the installer function.
import PHYS121
PHYS121.Installer()
```

```python
# Initialize Otter
import otter
grader = otter.Notebook("PHYS 121 - Pre-Lab 1.ipynb")
```

***
## Import Modules:
Execute the cell below to import a number of useful pre-built Python modules that will be used in the PHYS 121 pre-labs and labs.

```python
import numpy as np
import matplotlib.pyplot as plt
import pandas as pd
from matplotlib.pyplot import cm # used to generate a sequence of colours for plotting
from scipy.optimize import curve_fit
from IPython.display import HTML as html_print
from IPython.display import display, Markdown, Latex
import sys
```

```python
PHYS121.graderCheck(['PHYS121' in sys.modules, 'otter' in sys.modules, 'numpy' in sys.modules], ['PHYS121', 'otter', 'numpy'], grader.check('q0.1'))
```

***
## Introduction:

This pre-lab will be used to develop some insights that will help you begin to meet the learning objectives for Lab #1.  Furthermore, the pre-lab will introduce some of the Python functions that will be used to analyze the data that you collect during the lab.

When a quantity, such as the period of a pendulum, is measured many times using the same measurement technique, one obtains a **distribution** of measurements.  Most of the time (but not always), the data follow a **Gaussian** distribution which is also known as a **normal** distribution or **bell curve** (Fig. 1). The figure below shows an ideal Gaussian distribution which has been scaled (or normalized) so that the total area beneath the curve is one. The peak of the distribution occurs at the **mean** $\mu$ and its width is characterized by the **standard deviation** $\sigma$.  The horizonal axis in Fig. 1 shows the locations of $\mu$, $\mu\pm\sigma$, $\mu\pm 2\sigma$, and $\mu\pm 3\sigma$.

<p>
<center>
<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/a8fe7d061e8041f5f011dba0dfe43331d29d03b2/Pre-Lab-1/images/Standard_deviation_diagram_micro.svg" alt="image info" style="width: 85%; height: auto;" /><br>
<b>Fig. 1: An ideal Gaussian distribution.  Figure adapted from <a href="https://commons.wikimedia.org/wiki/File:Standard_deviation_diagram.svg">Wikimedia Commons</a>.</b></center>
</p>

When collecting experimental data, you tpyically won't see a smooth curve like the one in Fig. 1.  We often use **histograms** to display the distribution of experimental measurements.  First, we define ranges of values called **bins**, and then we count the number of datapoints that fall into each range.  For example, if you were to measure the speed of different objects in the asteroid belt, you could define your bins to be $1\rm\ km/s$ wide.  A measurement of $17.8\rm\ km/s$ (the average orbital speed of Ceres) would add a count to the $17-18\rm\ km/s$ bin.  

As more and more data is collected, histograms will approximate the shape of a Gaussian distribution. To exactly reproduce the ideal distribution in Fig. 1, one would need to make an infinite number of measurements of the quantity of interest and then plot a histogram of those data.  This is impractical, of course, so we can only hope to use a finite set of data to approximate an ideal distribution.  


**The 68-95-99.7 rule**<br>
The area between two points $x_1$ and $x_2$ beneath a normalized Gaussian distribution gives the probability that a single measurement $x$ will fall between those values.  For example, in Fig. 1 region shaded dark blue extends from $x_1=\mu-\sigma$ to $x_2=\mu+\sigma$ and has an area of $0.68$.  This observation implies that $68\%$ of measurements are expected to fall within one standard deviation $\sigma$ of the mean $\mu$.  In a similar way, $95\%$ of measurements are expected to fall within $\pm 2\sigma$ of the mean and $99.7\%$ of measurements are expected to fall within $\pm 3\sigma$ of the mean.

When estimating the uncertainty of a single measurement, one is trying to estimate the standard deviation that would be obtained if the measurement was repeated many times (without actually doing the repeated measurements). 

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 1 - Simulated Data ($N=20$)</span></div>

One of the advantages of using Jupyter notebooks for the lab manuals and pre-labs is that, when appropriate, it can be used to generate simulated data.  The simulated datasets can then be used to develop important insights that will help us interpret real experimental data.

The block of code below generates a set of $N$ simulated 'measurements' drawn from a Gaussian distribution with mean $\mu$ and standard deviation $\sigma$.  You do **not** need to understand the syntax of the code.  It is include only for the benefit of those interested in understanding the details of how the simulated data is generated.

```python
# Define values for mean & standard deviation of the Gaussian/normal distribution.
mu = 2.2 # mean 
sigma = 0.5 # standard deviation

# Enter the number of simulated data points to generate.
N = 20

# Generate the simulated data.
sim20 = np.random.normal(mu, sigma, N)

# Show simulated data using a Pandas dataframe.
df = pd.DataFrame(sim20, columns = ['data'])
df
```

In the cell below, you will use a function called 'Statistics' to produce a histogram of simulated data and calculate some basic statistics from your dataset (mean, standard deviation, and standard error). The 'Statistics' function is called as follows:
```python
import PHYS121
counts, centres, average20, stdDev20, stdError20, fig20 = PHYS121.Statistics(xData = sim20)
```
For additional information about the 'Statistics' function, see **Appendix A** near the end of this notebook.  

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 1.1</span>** **<span style="color:red">(1 mark)</span>**  

Execute the cell below to generate the required histogram.

*** Please do not change any of the variable names in the cell below used to generate the histogram ***

```python
# Plot a histogram of the simulated data and calculate some basic statistics.
counts, centres, average20, stdDev20, stdError20, fig20 = PHYS121.Statistics(data = sim20)
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 1.2</span>** **<span style="color:red">(1 mark)</span>**  

Comment on the shape of the histogram that was generated using the simulated data.  What features does it have in common with the ideal Gaussian distribution shown in Fig. 1?  What are the differences?  You will be graded on the completeness and clarity of your answer. 

***
**<span style="color:blue">Answer 1.2:</span>**

[//]: # (Please do not delete this comment or anything above it.  Anything below this comment can be deleted.)  

Double click this cell and enter your text here.  When done, hit 'Shift' + 'Enter' to execute the cell.  You may delete this text when entering your answer. 

<!-- END QUESTION -->

<!-- END QUESTION -->

***
The lines of code below are used to make a scatter plot of the simulated data.  The $x$-axis represents the trial number and  the $y$-axis is the corresponding value.  The dashed pink line marks the average value $\mu$, the solid red line marks $\mu+\sigma$, and the solid blue line marks $\mu-\sigma$.  

Again, you do **not** need to spend time understanding the syntax of the code.  It is merely provided for the benefit of those interested in seeing the specific details.  

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 1.3</span>** **<span style="color:red">(1 mark)</span>**  

Execute the cell below to generate the required scatter plot.

```python
# Define the size of the resulting figure (5" x 5")
fig = plt.figure(figsize=(5, 5), dpi=100)

# Plot the simulated data.  The 'ko' syntax specifies a plot of black (the 'k' in 'ko') circular data points (the 'o' in 'ko'). 
plt.plot(sim20, 'ko')

# Specify the range of the x-axis.
plt.xlim([-1, N + 1])

# Specify which ticks to label along the x-axis.
plt.xticks([0, 2, 4, 6, 8, 10, 12, 14, 16, 18, 20]);

#Label the x- and y-axes.
plt.xlabel('Trial number')
plt.ylabel('Simulated data')

# Plot a horizontal pink dashed line at mu
plt.plot([-1, N + 1], [average20, average20], color = 'pink', linestyle = '--')

# Plot a horizontal red line at mu + sigma
MuPlusSigma = average20 + stdDev20
plt.plot([-1, N + 1], [MuPlusSigma, MuPlusSigma], 'red')

# Plot a horizontal blue line at mu - sigma
MuMinusSigma = average20 - stdDev20
plt.plot([-1, N + 1], [MuMinusSigma, MuMinusSigma], 'blue');
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 1.4</span>**  **<span style="color:red">(1 mark)</span>**    

In the scatter plot above, the red line is placed at $\mu +\sigma$ along the $y$-axis and the blue line is placed at $\mu-\sigma$.  What fraction of points fall between the blue and red lines?  Enter the appropriate value for 'frac' in the cell below.  Delete the ... and replace it with a fraction or decimal number.  

*** Please do not change anything to the left of the equals sign. ***

```python
frac = ...
```

```python
PHYS121.graderCheck([frac], ['frac'], grader.check('q1.4'))
```

<!-- END QUESTION -->

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 2 - Simulated Data ($N=1000$)</span></div>

Here, we simply repeat everything from **Part 1** after increasing $N$ from $20$ to $1000$.  Using the same values of $\mu$ and $\sigma$ that were used in **Part 1**, generate the simulated data using the code block below.

```python
# Define values for mean & standard deviation of the Gaussian/normal distribution.
mu = 2.2 # mean 
sigma = 0.5 # standard deviation

# Enter the number of simulated data points to generate.
N = 1000

# Show all rows in the Panadas dataframe
pd.set_option('display.max_rows', 1000)

# Generate the simulated data.
sim1k = np.random.normal(mu, sigma, N)

# Show simulated data using a Pandas dataframe.
df = pd.DataFrame(sim1k, columns = ['data'])
df
```

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 2.1</span>** **<span style="color:red">(1 mark)</span>**  

Execute the cell below to generate the required histogram.

```python
# Plot a histogram of the simulated data and calculate some basic statistics.
counts, centres, average1k, stdDev1k, stdError1k, fig1k = PHYS121.Statistics(data = sim1k)
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 2.2</span>** **<span style="color:red">(1 mark)</span>**  

Execute the cell below to generate the required scatter plot.

```python
# Define the size of the resulting figure (5" x 5")
fig = plt.figure(figsize=(5, 5), dpi=100)

# Plot the simulated data.  The 'ko' syntax specifies a plot of black (the 'k' in 'ko') circular data points (the 'o' in 'ko'). 
plt.plot(sim1k, 'ko')

# Specify the range of the x-axis.
plt.xlim([-10, N + 10])

# Specify which ticks to label along the x-axis.
plt.xticks([0, 100, 200, 300, 400, 500, 600, 700, 800, 900, 1000]);

#Label the x- and y-axes.
plt.xlabel('Trial number')
plt.ylabel('Simulated data')

# Plot a horizontal pink dashed line at mu
plt.plot([-10, N + 10], [average1k, average1k], color = 'pink', linestyle = '--')

# Plot a horizontal red line at mu + sigma
MuPlusSigma = average1k + stdDev1k
plt.plot([-10, N + 10], [MuPlusSigma, MuPlusSigma], 'red')

# Plot a horizontal blue line at mu - sigma
MuMinusSigma = average1k - stdDev1k
plt.plot([-10, N + 10], [MuMinusSigma, MuMinusSigma], 'blue');
```

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 2.3</span>** **<span style="color:red">(1 mark)</span>**     

Compare the obtained values of $\mu$ and $\sigma$ from **Part 1** ($N = 20$) to those obtained in **Part 2** ($N = 1000$).  Are they very different or are they similar?  We changed $N$ by a factor of $1000/20 = 50$.  In comparison, by what factor did the standard deviation change?  What does this tell you about how the standard deviation $\sigma$ depends on the number of trials $N$?

***
**<span style="color:blue">Answer 2.3:</span>**

[//]: # (Please do not delete this comment or anything above it.  Anything below this comment can be deleted.)  

Double click this cell and enter your text here.  When done, hit 'Shift' + 'Enter' to execute the cell.  You may delete this text when entering your answer. 

<!-- END QUESTION -->

<!-- BEGIN QUESTION -->

***
**<span style="color:blue">Question 2.4</span>** **<span style="color:red">(1 mark)</span>**    

Compare the histograms obtained in **Part 1** ($N = 20$ trials) and **Part 2** ($N = 1000$ trials).  Which of the histograms allows you to better determine the position of the peak (or mean)?  Provide a clear justification for your answer.

***
**<span style="color:blue">Answer 2.4:</span>**

[//]: # (Please do not delete this comment or anything above it.  Anything below this comment can be deleted.)  

Double click this cell and enter your text here.  When done, hit 'Shift' + 'Enter' to execute the cell.  You may delete this text when entering your answer. 

<!-- END QUESTION -->

<!-- END QUESTION -->

***
From **Parts 1** and **2** of this pre-lab, you should have seen that the standard deviation $\sigma$, which characterizes the width of a Gaussian distribution, is independent of the number of trials $N$.  The standard deviation $\sigma$ is used to estimate the uncertainty of any single trial of an experiment.  The physical interpretation is that any single measurement has approximately a $68\%$ chance of falling within one standard deviation of the mean $\mu$.  That is, each measurement has a $68\%$ chance of falling between $\mu-\sigma$ and $\mu+\sigma$.

Although $\sigma$ is independent of $N$, our ability to locate the peak or mean value of a distribution of measurements improves as $N$ increases.  This is because the distribution of measurements better approximates an ideal Gaussian distribution with a smooth peak and symmetric tails as $N$ is increased.

The **standard error in the mean** (often referred to simply as the **standard error**) is an estimate of the uncertainty in the calculated mean from a set of data. We will denote the standard error by $\sigma_\mu$. The standard error is calculated by dividing the standard deviation $\sigma$ by the square root of the number of trials $N$:
\begin{align}
\sigma_\mu = \dfrac{\sigma}{\sqrt{N}}.
\tag{2}
\end{align}
Understanding where this expression comes from requires a proper treatment of the propagation of errors which is beyond the scope of PHYS 121.  However, if you choose to pursue a physics degree, we will revisit error propagation and standard error in the Modern Physics Laboratory course PHYS 232.

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 3 - Overlaying Histograms</span></div>

In this final part of the pre-lab, which won't be graded, we simply overlay the histograms of the simulated data generated in **Parts 1** and **2**.  To create the overlay of histograms, the 'HistOverlay()' function from PHYS121.py will be used.  For additional information about this function, see **Appendix B** near the end of this notebook.

```python
cnts, ctrs, fig2 = PHYS121.HistOverlay([sim20, sim1k])
```

In the plot above, you should see a pair of overlaid histograms.  The blue and orange bars represent the set of simulated data with $N=20$ and $N=1000$ trials, respectively.  You should notice that the orange data better approximate an ideal Gaussian distribution and thus the mean value of the distribution is easier to locate.  This plot also emphasizes that the two distributions have approximately the same width, related to the fact that the standard deviation $\sigma$ is independent of the size $N$ of the dataset.

To summarize, the two important observations that can be made from this analysis are:
- The standard deviation $\sigma$ is independent of the number of trials $N$.
- Datasets with a large number of trials result in smoother distributions with allows for a more precise determination of the mean $\mu$.

If both of the observations above are true, then $\sigma$ **cannot** be a good estimate of the uncertainty in the mean or average value of a set of $N$ measurements.  Instead, the standard deviation $\sigma$ represents the uncertainty in the value of any single trial of the experiment.  It is the standard error $\sigma_\mu=\sigma/\sqrt{N}$ that represents the uncertainty in the mean or average value of a set of $N$ measurements.

<!-- BEGIN QUESTION -->

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 4 - Feedback and Submission</span></div>

**<span style="color:blue">Question 4.1:</span>**  

We welcome your feedback on the PHYS 121 pre-labs!  Please feel free to include any comments you have about this pre-lab in the cell below.  Your comments will be taken into consideration when revising/improving the PHYS 121 labs and pre-labs.  You can suggest improvements, point out anything that was unclear, comment on the strengths and weaknesses of the pre-lab, ...

This question is optional and will have no impact on your pre-lab or lab grade.

***
**<span style="color:blue">Answer 4.1:</span>**

[//]: # (Please do not delete this comment or anything above it.  Anything below this comment can be deleted.)  

Double click this cell and enter your text here.  When done, hit 'Shift' + 'Enter' to execute the cell.  You may delete this text when entering your answer. 

<!-- END QUESTION -->

***
Once you've completed this notebook:
- Save your work.
- Run 'grader.check_all()' to confirm that you've completed all required tasks.
- Run 'grader.export()' to generate a .zip file containing all of the materials that you will submit.
- Download the generated .zip file using the link that appears **below** the ```grader.check_all(...)``` cell.
- $\Leftarrow$ Please do **NOT** download the zip file from the frame along the left-hand side of the screen. 
- Upload the .zip file to the PHYS 121 Canvas gradebook.
- **Do NOT change the name of the .zip file.**
- **Do NOT modify the contents of the .zip file.**

Here is a <a href = "https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/general/gifs/Submission.gif">GIF</a> showing how these steps are completed.  Once your completed notebook has been uploaded to the Canvas gradebook, you're done!

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
display(grader.check_all())
PHYS121.saveMessage()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
PHYS121.waitMessage()
grader.export(files=["PHYS121_DataLogger.txt"])
PHYS121.proceedMessage()
```

# <div style="background-color:rgba(255, 204, 255, 0.5); text-align:center; vertical-align: middle; padding:40px 0; margin-top:30px"><span style="color:rgba(102, 0, 204, 1);">Part 5 - Playground (optional)</span></div>

Feel free to add as many cells as you like below and use them as a playground for further independent investigations.  These cells won't be graded, so feel free to use them in any way that you like.  For example, you could compare Gaussian distributions with different standard deviations and/or means. 

```python
# Here's an empty code cell that you can use.
```

```python
# Here's another empty code cell that you can use.
```

```python
# Here's yet another empty code cell that you can use.  
# If you need more, you can add cells using the '+' icon in the menu bar at to the top of the screen.
```

### <div style="background-color:rgba(255, 255, 102, 0.5); text-align:left; padding:20px 0; margin-top:20px">$\quad$Appendix A &ndash; The Statistics Function...</div>

The statistics function is called as follows:
```python
Statistics(data, nbins = 10, xlabel = 'x-axis', xUnits = '', normalized = False)
```
The data input is required, all other arguments are optional with default values set.  The function returns the following outputs:
```python
counts, centres, average, stdDev, stdError, fig
```
The first two outputs are lists, the next three are numerical values, and 'fig' is the formatted plot. The counts list gives the number of counts in each bin of the histogram and the centres list gives the position of the centre of each bin.

The function calculates and reports the mean, standard deviation, and standard error of the provided data.  It also plots a histogram of the data.  

### Statistics Example Implmentation
The code block below shows an implementation of 'Statistics'.
```python
import PHYS121
theta20a = [2.02, 1.93, 1.92, 1.96, 2.03, 2.03, 1.96, 2.03, 2.06, 2, 2.03, 2.12, 2.07, 1.99, 1.99, 1.95, 2.03, 2.12, 2.03, 2.09, 2.03, 2.03, 2.01, 2.04, 2.03, 2.04, 1.99, 1.99, 1.97, 1.98]
cnts, ctrs, mean, stdDev, stdError, fig20a = PHYS121.Statistics(theta20a, 12, 'Period', 's', True);
```

If you're interested in generating your own histograms with customized formatting, see the following Python-based tutorial: https://cmps-people.ok.ubc.ca/jbobowsk/Python/html/Jupyter%20Histogram.html.
***

### <div style="background-color:rgba(255, 255, 102, 0.5); text-align:left; padding:20px 0; margin-top:20px">$\quad$Appendix B &ndash; The HistOverlay Function...</div>

The HistOverlay function is called as follows:
```python
HistOverlay(dataArray, nbins = 10, xlabel = 'x-axis', xUnits = '',  normalized = True, transparency = 0.75)
```
The dataArray (a list of lists) input is required, all other arguments are optional with default values set.  The function returns the following outputs:
```python
countsArray, centresArray, fig
```
The first two outputs are lists of lists and 'fig' is the formatted plot. 

The dataArray is of the form dataArray = [[dataset1], [dataset2], ... [datasetN]].  The function plots the histograms of each of the datasets on a single graph.  If the normalized argument is True, then histogram is scaled such that the area under the distributions is 1.  This is a good option when comparing two different distributions.  The transparency argument (a number between 0 and 1) sets the transparency of each of the individual histograms.

### HistOverlay Example Implementation
The code block below shows an implementation of 'HistOverlay'.
```python
import PHYS121
cnts, ctrs, fig = PHYS121.HistOverlay([theta20a, theta20], 8, 'Period', 's', True, 0.75)
```
***

<img src="https://raw.githubusercontent.com/UBC-Okanagan-Physics-Labs/PHYS-121-images/main/general/images/ubc-logo-full.jpg" width="500"/>

Last update: January 18, 2025

Tests for q0.1 (points: None)
```
>>> 'PHYS121' in sys.modules and 'otter' in sys.modules and ('numpy' in sys.modules)
True
```
failure_message: Make sure you execute the cells containing "import PHYS121" and "import otter" as well as the cell beneath the heading "Importing Packages/Modules".


Tests for q1.4 (points: None)
```
>>> isinstance(frac, float)
True
```
failure_message: frac must be a floating-point number (i.e. a fraction such as 13/20 or a decimal number such as 0.65).
```
>>> frac > 0 and frac < 1
True
```
failure_message: frac should be a fraction or decimal number between 0 and 1.
